<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Programmierung für KI
### Winterersemester 2025/26
Prof. Dr. Heiner Giefers

**Aufgabe 1: Schreiben Sie ein Programm, das wiederholt Zahlen einliest, bis der Benutzer den Befehl `done` eingibt. Sobald `done` eingegeben wurde, geben Sie die Summe, die Anzahl und den Durchschnitt der Zahlen aus. Wenn der Benutzer etwas anderes als eine Zahl eingibt, erkennen Sie seinen Fehler mit `try` und `except` und verlassen die Funktion, indem Sie die Fehlermeldung `Es werden nur Zahlen akzeptiert. Beenden mit 'done'` als Zeichenkette (mit `return`) 
zurückgeben.**

*Erwartete Ausgabe:*

    Bitte eine Zahl eingeben: 2
    Bitte eine Zahl eingeben: 4
    Bitte eine Zahl eingeben: 6
    Bitte eine Zahl eingeben: done
    Anzahl: 3 Summe: 12.0 Durchschnitt: 4.0

<!-- grade_id: cell-aaa380245b05e7d3, points: 0, grade: false, solution: true -->
```python
def durchschnitt():
    anzahl = 0
    summe = 0
    durchschnitt = 0
    end = False

    while end != True:
        eingabe = input("Bitte eine Zahl eingeben:")
        
        if eingabe == "done":
            end = True
            erg = "Anzahl: " + str(anzahl) + " Summe: " + str(summe) + " Durchschnitt: " + str(durchschnitt)
            return erg
        
        try:
            zahl = float(eingabe)        
            anzahl = anzahl + 1
            summe = summe + zahl
            durchschnitt = summe / anzahl
        except ValueError:
            return "Es werden nur Zahlen akzeptiert. Beenden mit 'done'"
          
    return erg
    
```

<!-- grade_id: cell-09543649eed06a8a, points: 2, grade: true, solution: false -->
```python
"""Teste die Funktion durchschnitt auf Gültigkeit."""

def myinput(inputs):
    def foo(a):
        print(a)
        input.counter += 1
        if input.counter > len(inputs) :
            raise "Nicht genügend inputs"
        else:
            return inputs[input.counter-1]
    foo.counter = 0
    return foo

from nose.tools import assert_equal

input_old = input
# Erster Testlauf, Abschluss der Eingabe mit "done"
input = myinput(['2','4', '6', 'done'])
s1 = durchschnitt()
# Zweiter Testlauf, Abschluss der Eingabe mit fehlerhaftem Input
input = myinput(['2','4','falsch'])
s2 = durchschnitt()
input = input_old

assert "Anzahl: 3 Summe: 12.0 Durchschnitt: 4.0" in s1, "Die Ausgabe stimmt nicht"
assert "Es werden nur Zahlen akzeptiert" in s2, "Die Fehlermeldung ist nicht korrekt"
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
Cell In[6], line 14
     11     foo.counter = 0
     12     return foo
---> 14 from nose.tools import assert_equal
     16 input_old = input
     17 # Erster Testlauf, Abschluss der Eingabe mit "done"

ModuleNotFoundError: No module named 'nose'
```

**Aufgabe 2: Schreiben Sie ein weiteres Programm, das eine Liste von Zahlen wie oben abfragt und am Ende sowohl das Maximum als auch das Minimum der Zahlen anstelle des Durchschnitts ausgibt.**

<!-- grade_id: cell-684d8a49ce7ad30e, points: 0, grade: false, solution: true -->
```python
def minmax():
    anzahl = 0
    minimum = None
    maximum = None

    
    while True:
        eingabe = input("Bitte eine Zahl eingeben:")
            
        if eingabe == "done":
            erg = "Minimum: " + str(minimum) + " Maximum: " + str(maximum)
            return erg
            
        try:
            zahl = float(eingabe)        
            if minimum == None or minimum > zahl:
                minimum = zahl
                
            if maximum == None or maximum < zahl:
                maximum = zahl
        except ValueError:
            return "Es werden nur Zahlen akzeptiert. Beenden mit 'done'"
              
    return erg
```

<!-- grade_id: cell-806efb201458eefd, points: 2, grade: true, solution: false -->
```python
"""Teste die Funktion minmax auf Gültigkeit."""

def myinput(inputs):
    def foo(a):
        print(a)
        input.counter += 1
        if input.counter > len(inputs) :
            raise "Nicht genügend inputs"
        else:
            return inputs[input.counter-1]
    foo.counter = 0
    return foo

from nose.tools import assert_equal

input_old = input
input = myinput(['2','4','1','done'])
s1 = minmax()
input = myinput(['2','4','test'])
s2 = minmax()
input = input_old

assert "Minimum: 1.0 Maximum: 4.0" in s1, "Die Ausgabe stimmt nicht"
assert "Es werden nur Zahlen akzeptiert" in s2, "Die Fehlermeldung ist nicht korrekt"
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
Cell In[4], line 14
     11     foo.counter = 0
     12     return foo
---> 14 from nose.tools import assert_equal
     16 input_old = input
     17 input = myinput(['2','4','1','done'])

ModuleNotFoundError: No module named 'nose'
```

```python

```