<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Programmierung für KI
### Winterersemester 2025/26
Prof. Dr. Heiner Giefers

**Aufgabe 1: Schreiben Sie eine while-Schleife, die beim letzten Zeichen in der Zeichenkette beginnt und sich rückwärts bis zum ersten Zeichen in der Zeichenkette vorarbeitet, wobei jeder Buchstabe in einer eigenen Zeile ausgegeben wird, natürlich rückwärts.**

<!-- grade_id: cell-9f6dc7ff3dbe0557, points: 0, grade: false, solution: true -->
```python
def drucke_zeichen_rueckwaerts(zeichenkette):

  index = len(zeichenkette) - 1
  while index >= 0:
    # Das Zeichen direkt ausgeben. print() fügt den Zeilenumbruch hinzu.
    print(zeichenkette[index])
    index -= 1
```

```python
%%capture cap --no-stderr --no-display
drucke_zeichen_rueckwaerts("Hello")
```

<!-- grade_id: cell-a2fe2825342aaa23, points: 1, grade: true, solution: false -->
```python
"""Test der extrahiereZeichen Funktion"""
from nose.tools import assert_equal
assert_equal(str(cap), 'o\nl\nl\ne\nH\n', "Die Ausgabe stimmt nicht")
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
Cell In[3], line 2
      1 """Test der extrahiereZeichen Funktion"""
----> 2 from nose.tools import assert_equal
      3 assert_equal(str(cap), 'o\nl\nl\ne\nH\n', "Die Ausgabe stimmt nicht")

ModuleNotFoundError: No module named 'nose'
```

**Aufgabe 2: Nehmen Sie den folgenden Python-Code, der eine Zeichenkette speichert:**

**`str = 'X-DSPAM-Confidence:0.8475'`**

**Verwenden Sie `find` und String-Slicing, um den Teil der Zeichenkette nach dem Doppelpunkt zu extrahieren, und verwenden Sie dann die Funktion `float`, um die extrahierte Teilzeichenkette in eine Fließkommazahl zu konvertieren.**

<!-- grade_id: cell-997127479ac5692f, points: 0, grade: false, solution: true -->
```python
def extrahiere_zeichen(zeichenkette):
    index = zeichenkette.find(":")
    rueck = float(zeichenkette[index+1:])
    return rueck
```

<!-- grade_id: cell-a2390f7eb31779fa, points: 1, grade: true, solution: false -->
```python
"""Test der extrahiereZeichen Funktion"""
from nose.tools import assert_equal
assert_equal(0.8475, extrahiere_zeichen("X-DSPAM-Confidence:0.8475"))
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
Cell In[5], line 2
      1 """Test der extrahiereZeichen Funktion"""
----> 2 from nose.tools import assert_equal
      3 assert_equal(0.8475, extrahiere_zeichen("X-DSPAM-Confidence:0.8475"))

ModuleNotFoundError: No module named 'nose'
```

**Aufgabe 3: Schreiben Sie eine Funktion `count_vocals`, welche die Vokale in einem als Parameter übergebenen Text zählt und die Summe aller Vokale zurück gibt.**

<!-- grade_id: cell-6405df4aa3f9f2b7, points: 0, grade: false, solution: true -->
```python
def count_vocals(text):
    count = 0
    vocals = "aeiou"
    
    # Loop through each character of the input text
    for char in text.lower():
        # Check if the character is one of the vocals
        if char in vocals:
            count += 1
            
    return count
```

<!-- grade_id: cell-8012e60ca6ad29a0, points: 1, grade: true, solution: false -->
```python
"""Test der countVocals Funktion"""
from nose.tools import assert_equal
# Beispiel - Text Quelle: https://de.wikipedia.org/wiki/Vokal
assert_equal(58, count_vocals("ein vokal oder selbstlaut ist ein sprachlaut, bei dessen artikulation der phonationsstrom weitgehend ungehindert durch den mund ausströmen kann. vokale sind im allgemeinen stimmhaft."))
assert_equal(58, count_vocals("Ein Vokal oder Selbstlaut ist ein Sprachlaut, bei dessen Artikulation der Phonationsstrom weitgehend ungehindert durch den Mund ausströmen kann. Vokale sind im Allgemeinen stimmhaft."))
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
Cell In[7], line 2
      1 """Test der countVocals Funktion"""
----> 2 from nose.tools import assert_equal
      3 # Beispiel - Text Quelle: https://de.wikipedia.org/wiki/Vokal
      4 assert_equal(58, count_vocals("ein vokal oder selbstlaut ist ein sprachlaut, bei dessen artikulation der phonationsstrom weitgehend ungehindert durch den mund ausströmen kann. vokale sind im allgemeinen stimmhaft."))

ModuleNotFoundError: No module named 'nose'
```