<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Programmierung für KI
### Winterersemester 2025/26
Prof. Dr. Heiner Giefers

**Übung 1: Schreiben Sie eine Funktion `lese_datei`, die eine Datei einliest und den Inhalt der Datei in Großbuchstaben (Zeile für Zeile) ausgibt.**

<!-- grade_id: cell-c1e55762bc63dc93, points: 0, grade: false, solution: true -->
```python
def lese_datei(filename):
    fhand = open(filename)
    big_fhand = fhand.read()
    return big_fhand.upper()
```

```python
%%capture cap --no-stderr --no-display
with open('myfile.txt', "w") as myfile:
    myfile.write("Ein neuer\nText.\nEnde")
lese_datei("myfile.txt")
```

Output:
```
'EIN NEUER\nTEXT.\nENDE'
```

<!-- grade_id: cell-0aa5a4884c58af6f, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal('EIN NEUER\n\nTEXT.\n\nENDE\n', str(cap))
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
Cell In[3], line 1
----> 1 from nose.tools import assert_equal
      2 assert_equal('EIN NEUER\n\nTEXT.\n\nENDE\n', str(cap))

ModuleNotFoundError: No module named 'nose'
```

**Übung 2: Schreiben Sie eine Funktion `berechne_durchschnitt`, die eine Datei einliest und in der Datei nach Zeilen der Form 
`X-DSPAM-Confidence: 0.8475` sucht. Die Zahl kann sich dabei ändern, die anderen Zeichen nicht. Insbesondere muss die Zeile immer mit `X-DSPAM...` beginnen.<br>Extrahieren Sie die Fließkommerzahl indem Sie die Zeile auftrennen. Berechnen Sie die Summe der Fließkommerzahlen und geben Sie den durchschnittlichen Wert der Fließkommerzahl zurück.**

<!-- grade_id: cell-8e47dfd8aabc8a20, points: 0, grade: false, solution: true -->
```python
def berechne_durchschnitt(dateiname):
    summe_werte = 0.0
    anzahl_zeilen = 0
    with open(dateiname, 'r') as datei:
        for zeile in datei:
            zeile = zeile.strip()
            if zeile.startswith("X-DSPAM-Confidence:"):
                doppelpunkt_pos = zeile.find(':')
             
                wert_str = zeile[doppelpunkt_pos + 1:].strip()
                wert = float(wert_str)
                summe_werte += wert
                anzahl_zeilen += 1

    if anzahl_zeilen > 0:
        durchschnitt = summe_werte / anzahl_zeilen
        return durchschnitt
    else:
        return 0.0
```

<!-- grade_id: cell-4ead248e87efeb5f, points: 2, grade: true, solution: false -->
```python
from nose.tools import assert_almost_equal


from urllib.request import urlopen
import os.path

fname = "mbox-short.txt"
if not os.path.exists(fname):
    url = "https://raw.githubusercontent.com/fh-swf-hgi/py4e/master/code3/mbox-short.txt"
    with urlopen(url) as response:
        content = response.read().decode('utf-8')
        with open(fname, "w") as file:
            file.write(content)



# Durchschnittswert prüfen
assert_almost_equal(0.7507185185185187, berechne_durchschnitt(fname))
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
Cell In[5], line 1
----> 1 from nose.tools import assert_almost_equal
      4 from urllib.request import urlopen
      5 import os.path

ModuleNotFoundError: No module named 'nose'
```